---
title: "Tidyverse Tutorial: R4DS Ch 3 & 5"

output: 
  learnr::tutorial:
    css: css/style.css
    theme: cerulean
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr)
knitr::opts_chunk$set(echo = FALSE)
library(tidyverse)
library(nycflights13)
```

###### This tutorial was written by the following students at Ritsumeikan Asia Pacific University in Japan as part of a seminar project under the supervision of Prof. William Claster. To collaborate please contact us.

<!-- ![picture](images/MC4ML.JPG) <img src="images/MC4ML.JPG" alt="drawing" style="width:80px;" /> -->


<!-- [![RProgramming](./images/MC4ML.JPG style="width:80px;" "Mathematics and Programming for Machine Learning with R")](https://www.amazon.com/Mathematics-Programming-Machine-Learning-R/dp/0367507854/ref=sr_1_1?crid=7Y7JVLAZ72TE&dchild=1&keywords=william+claster&qid=1630329202&sprefix=william+Claster%2Caps%2C358&sr=8-1) -->
<font size="2">

###### Kamal Al-Tway (KA), Dang Ngo (DN), Soujanya Dash (SD), Yuji Yamaguchi (YY)



###### To learn to program machine learning algorithms from scratch in R consider William Claster's [Mathematics and Programming for Machine Learning with R](https://www.amazon.com/Mathematics-Programming-Machine-Learning-R/dp/0367507854/ref=sr_1_1?crid=7Y7JVLAZ72TE&dchild=1&keywords=william+claster&qid=1630329202&sprefix=william+Claster%2Caps%2C358&sr=8-1)  
</font> 

####### Do as many of the exercises as you like

## Section 3.2.4

### 3.2.4 Easy SD

 

After installing Tidyverse package, given the mpg dataset, plot a scatterplot between any two continuous variable. If you don't recall the variables you can type "?mpg" (without the quotation marks) or use colnames(mpg)


```{r two-plus-two, exercise=TRUE}

```

```{r two-plus-two-hint}
ggplot(data = ) + 
  geom_point(mapping = aes(x = , y = ))
```




### 3.2.4 Intermediate SD


Modify the following code so that the legends of the aesthetics is not shown:

```{r print-limit, exercise=TRUE, exercise.eval=TRUE}

x <- ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy, color = class))
print(x)
```

```{r print-limit-hint}
use the show.legend parameter of the geo_point() function.
```

```{r print-limit-solution}
x <- ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy, color = class),show.legend = FALSE)
print(x)
```

```{r print-limit-check}
gradethis::grade_this_code()
```


## Section 3.3.1

### 3.3.1 Easy DN


Write the R code for putting mpg dataset into a mpg dataframe

```{r D1, exercise=TRUE}
mpg<- data.frame()
```

```{r D1-solution, include=FALSE}
mpg<- data.frame(mpg)
```

```{r D1-check}
gradethis::grade_this_code()
```


Write the R code for Viewing the mpg tibble
```{r D2, exercise=TRUE}

```
```{r D2-hint}
"use the View() function"
```

```{r D2-solution}
View(mpg)
```

```{r D2-check}
gradethis::grade_this_code()
```




### 3.3.1 Easy KA

Write the code that will load tidyverse:

```{r 22K, exercise=TRUE}

```
```{r 22K-hint-1}
"first part is library()"

```
```{r 22k-hint-2}
"you need to insert the word tidyverse"
```

```{r 22K-solution}
library(tidyverse)
```
```{r 22K-check}
gradethis::grade_code()
```

### 3.3.1 Easy KA

If you want to learn more about the mpg tibble what code could you run?

```{r two-plus-three, exercise=TRUE}

```
```{r two-plus-three-hint-1}
"use the ? sign"
```

```{r two-plus-three-solution}
?mpg
```

```{r two-plus-three-check}
gradethis:: grade_code()
```



### 3.3.1 Easy SD


Edit the code by changing the color aesthetics to alpha aesthetics, mapping it to the class attribute:
**HINT**: *print x to view the graph*

```{r add-function, exercise=TRUE, exercise.lines = 5}
x <- ggplot(data = mpg) + geom_point(mapping = aes(x = displ, y = hwy, color = red)) 
print()
```

```{r add-function-hint}
 # alpha is a characteristic of the geometric element and determines the opacity.
```

```{r add-function-solution}
x <- ggplot(data = mpg) + geom_point(mapping = aes(x = displ, y = hwy, alpha = class)) 
print(x)
```

```{r add-function-check}
gradethis::grade_this_code()
```



###  3.3.1 Intermediate DN

Refer to the mpg data set.  

Plot a geom_point graph with displ on x-axis, hwy on y-axis, and map the shape of points to drv
```{r D12, exercise=TRUE}
ggplot(    ) + geom_point(    )
```

```{r D12-hint}
"Mapping the shape of points to drv means that the argument should be included inside aesthetic"
```

```{r D12-solution}
ggplot(data = mpg) + geom_point(mapping = aes(x = displ, y = hwy, shape= drv))
```

```{r D12-check}
gradethis::grade_this_code()
```




```{r D13}
quiz(
  question("Which of the following are considered asethetics in geom_point ?",
    answer("fill", correct = TRUE),
    answer("shape", correct = TRUE),
    answer("size", correct= TRUE),
    answer("length")
  )
)
```

### 3.3.1 Intermediate KA

<!-- ![picture](images/Rplot22.png) -->
<img src="images/Rplot22.png" alt="drawing" style="width:400px;"/>

enter the code that creates the above graph

```{r two-plus-four, exercise=TRUE}

```

```{r two-plus-four-hint-1}
"you need to use both ggplot() and ggplot_point()"
```

```{r two-plus-four-hint-2}
"ggplot(data = mpg) + 
  geom_point(mapping = aes(x= , y= ))"
```

```{r two-plus-four-solution}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy))
```

```{r two-plus-four-check}
gradethis:: grade_code()
```

### 3.3.1 Intermediate KA

<!-- ![picture](images/Rplot23.png) -->
<img src="images/Rplot23.png" alt="drawing" style="width:400px;"/>

Although using size for a discrete variable is not advised, nevertheless, complete the code to create the above graph

```{r two-plus-five, exercise=TRUE}

ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy, ))
             
```

```{r two-plus-five-hint-1}
"missing part is size"
```

```{r two-plus-five-solution}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy, size = class))
```
```{r two-plus-five-check}
gradethis:: grade_code()
```

### 3.3.1 Intermediate SD



Question to verify that one has understood the basics of ggplot:

```{r quiz7}
quiz(
  question("Which of the following modifies a ggplot or theme object by adding on new components?",
    answer("+",correct = TRUE),
    answer("*"),
    answer("/"),
    answer(".")
  )
)
```


```{r quiz8}
quiz(
  question("What are the basic units in the grammar of graphics?",
    answer("Data and layers"),
    answer("Geometries and data"),
    answer("The data, the aesthetics, and the geometries", correct = TRUE),
    answer("Aesthetics and geometries")
  )
)
```




### 3.3.1 Advanced DN  


Divide mpg dataframe into 2 halves and plot 2 separate geom_point graphs  with displ on x-axis and hwy on y-axis  respectively.
**Note**:"spliting the mpg dataframe requires floor(), set.seed(), sample(), seq(). You many want to get help on this functions first". Also, note we are getting a random sample for each of the 2 halves.

```{r D23, exercise=TRUE}
half_size <- floor(0.5*  )
set.seed(123123)
mpg_half <- sample(seq(   ), size= half_size)

mpg_1 <- mpg[mpg_half, ]
mpg_2 <- mpg[-mpg_half, ]

ggplot(data= mpg_1) +
  geom_point(  )
ggplot(data=mpg_2) +
  geom_point(  )
```

```{r D23-solution}
half_size <- floor(0.5*nrow(mpg))
set.seed(123123)
mpg_half <- sample(seq(nrow(mpg)),size= half_size)

mpg_1 <- mpg[mpg_half, ]
mpg_2 <- mpg[-mpg_half, ]

ggplot(data= mpg_1) +
  geom_point(mapping = aes(x= displ, y = hwy))
ggplot(data=mpg_2) +
  geom_point(mapping = aes(x= displ, y = hwy))
```

```{r D23-check}
gradethis::grade_this_code()
```

### 3.3.1 Challenging KA

<!-- ![picture](images/Rplot24.png) -->
<img src="images/Rplot24.png" alt="drawing" style="width:400px;"/>


run the code that creates the graphs above

```{r two-plus-six, exercise=TRUE}

```
```{r two-plus-six-hint-1}
"you need to add shape = class"
```
```{r two-plus-six-solution}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy, shape = class))
```
```{r two-plus-six-check}
gradethis:: grade_code()
```

### 3.3.1 Challenging  SD

Make a 3D scatterplot to analyse the relation between sepal length, sepal width and petal length.
```{r Challenge, exercise=TRUE,exercise.eval = TRUE}
#install.packages("plot3D")
library(plot3D);library(tidyverse)
```

```{r Challenge-hint}
#Step1 run data("iris") & head(iris)
#Step2 Define x,y, and z coordinates. For example, x<-iris$Sepal.Length
#Step3 use scatter3D(x,y,z,clab = c("Sepal","Width(cm)"))
#See this link for more info: https://preludeinr.com/graphs-and-plots/fancy-a-little-bit-of-3d/
```
## Section 3.5.1 

### 3.5.1 Easy SD


Let's try making facets for our mpg data and displ as x and hwy as y. Use facet_wrap() on any one categorical variable & facet_grid() on any two categorical variables. Can you see the difference?
```{r Facet, exercise=TRUE}
ggplot(data=)+geom_point()+
  
ggplot(data=)+geom_point()+  
  
```

```{r Facet-hint-1}
"the first argument of facet_wrap() is a tilda followed by the categorical variable "
"remember a tilda looks like this: ~"
```
```{r Facet-hint-2}
"the argument of facet_grid() has two variables separated by a tilda "
```

```{r Facet-solution}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy)) + 
  facet_wrap(~ class, nrow = 2)
  
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy)) + 
  facet_grid(drv ~ cyl)
```



### 3.5.1 Intermediate SD

*Where do we use geom_smooth? A linear model can be created on a scatterplot by using this technique which will deliver the 'best fit' line or a line average of all the points.*

For the airquality dataset, given the scatterplot, find the 'best fit' line:

```{r airquality, exercise=TRUE, exercise.eval=TRUE}
library(tidyverse)

ggplot(data = airquality,aes(y=Wind,x=Temp))+geom_point()
```

```{r airquality-hint}
geom_point()+geom_smooth()
```
```{r airquality-solution}
ggplot(data = airquality,aes(y=Wind,x=Temp))+geom_point()+geom_smooth()
```

```{r airquality-check}
gradethis::grade_this_code()
```

<!-- ### 3.5.1 Challenging  SD -->

<!-- Let's do some visualization with the basic knowledge. We are working with the mpg dataset. Questions will be in steps: -->

<!-- Step 1: Find the number of rows and columns of the dataset. -->

<!-- Step 2: Find the table of manufacturers. You want to show each how many instances (rows) in our dataset are there for each possible manufacturer. It would look something like this: -->

<!-- <!-- ![picture](images/351SDChalleng.JPG) --> -->
<!-- <img src="images/351SDChalleng.JPG" alt="drawing" style="width:400px;"/> -->


<!-- Step 3: Make a bar chart to represent the number of manufacturers. -->

<!-- ```{r Step1, exercise=TRUE, exercise.eval=TRUE} -->

<!-- ``` -->

<!-- ```{r Step1-hint1} -->
<!-- # Use table() and then index the attribute you want to show ($) -->
<!-- ``` -->

<!-- ```{r Step1-hint2} -->
<!-- # geom can take point, smooth, bar etc. -->
<!-- ``` -->

<!-- ```{r Step1-solution} -->
<!-- nrow(mpg) -->
<!-- ncol(mpg) -->
<!-- table(mpg$manufacturer) -->
<!-- qplot(manufacturer, data=mpg, geom="bar", fill=manufacturer) -->
<!-- ``` -->

<!-- # ```{r step1-check} -->
<!-- # gradethis::grade_this_code() -->
<!-- # ``` -->



########################################
## Section 3.6.1 

### 3.6.1 Easy DN



```{r D3}
quiz(
  question("Which of the following are considered variables in facet_wrap ?",
    answer("rows"),
    answer("cols"),
    answer("nrow", correct= TRUE),
    answer("ncol", correct= TRUE)
  )
)
```



```{r D4}
quiz(
  question("Which of the following are considered variables in facet_grid ?",
    answer("rows", correct= TRUE),
    answer("cols", correct= TRUE),
    answer("nrow"),
    answer("ncol")
  )
)
```

### Easy 3.6.1 KA
```{r quiz1a}
question("which best describes what a geom is?",
  answer("tells the plot how to display the data, for example as points, bars, line(s)", correct = TRUE),
  answer("tells the plot the the colors, sizes, and shapes to use "),
  answer("focuses on how individual elements of data are displayed")
)
```

### 3.6.1 Easy KA
<!-- ![picture](images/Rplot28.png) -->
<img src="images/Rplot28.png" alt="drawing" style="width:400px;"/>


run the code to create the above graph using ggplot()


```{r two-plus-eight, exercise=TRUE}

```

```{r two-plus-eight-solution}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy))
```

```{r two-plus-eight-check}
gradethis:: grade_code()
```



#### 3.6.1 Easy YY
Change the number of nrow in the code below and run the code to understand what the meaning of nrow is.
```{r 136, exercise=TRUE, exercise.eval=TRUE}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy)) + 
  facet_wrap(~ class, nrow = 2)
```
```{r 136-hint-1}
"nrow means number of the rows "
```


```{r quiz236}
question("What does geom_smooth do?",
  answer("Eliminates jitter in the data"),
  answer("Aids the eye in seeing a pattern in the data", correct = TRUE),
  answer("Reverses the x and y axes"),
  answer("None of these")
)
```


### 3.6.1 Easy SD

*Labeling is important for both the analyst and for those who use those visualizations*

Let's label the facets we made by adding a title of "Facet wrap" with the lab() function
```{r Labeling, exercise=TRUE}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy)) + 
  facet_wrap(~ class, nrow = 2)
```

```{r Labeling-hint}
# Use labs()
```

```{r Labeling-solution}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy)) + 
  facet_wrap(~ class, nrow = 2)
labs(
  title = "Facet wrap",
)
```

```{r Labeling-check}
gradethis::grade_this_code()
```



### 3.6.1 Intermediate DN


In spite of the fact that the graph below may actually obscure the relationship between the displ and hwy, as an exercise recreate the R code necessary to generate the following graph.
```{r D14-graph}
ggplot(data=mpg, mapping = aes(x= displ, y= hwy)) +
  geom_point() +
  geom_line()
```


```{r D14, exercise=TRUE}
ggplot(data=mpg, mapping = aes(x= displ, y= hwy )) 
  
```

```{r D14-hint}
"You will need 2 geom elements"
```

```{r D14-solution}
ggplot(data=mpg, mapping = aes(x= displ, y= hwy)) +
  geom_point() +
  geom_line()
```

```{r D14-check}
gradethis::grade_this_code()
```
```{r quiz361DN}
question("Instead of connecting every point as was shown above, what might better show the relationship between the two variables displ and hwy?",
  answer("Replace the dots with bars (using geom_bar())"),
  answer("add color to aes() to reflect the manufacturer"),
  answer("Add in a geom_smooth()", correct = TRUE),
  answer("None of these")
)
```


Recreate the R code necessary to generate the following graphs
```{r D15-graph}
ggplot() +
  geom_point(data=mpg, mapping= aes(x=displ, y= hwy, color= drv)) +
  geom_line(data=mpg, mapping = aes(x=displ, y= hwy, color= drv))

```

```{r D15, exercise=TRUE}
ggplot() +
  geom_point() +
  geom_line()

```

```{r D15-hint}
"You should map the color of points to drv variable"
```

```{r D15-solution}
ggplot() +
  geom_point(data=mpg, mapping= aes(x=displ, y= hwy, color= drv)) +
  geom_line(data=mpg, mapping = aes(x=displ, y= hwy, color= drv))
```

```{r D15-check}
gradethis::grade_this_code()
```  

### 3.6.1 Intermediate KA

<!-- ![picture](images/Rplot29.png) -->
<img src="images/Rplot29.png" alt="drawing" style="width:400px;"/>

fix the code to create the graph above

```{r two-plus-nine, exercise=TRUE}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy, linetype = drv))
```
```{r two-plus-nine-hint-1}
"try geom _smooth() rather than geom_point()"
```

```{r two-plus-nine-solution}
ggplot(data = mpg) + 
  geom_smooth(mapping = aes(x = displ, y = hwy, linetype = drv))
```

```{r two-plus-nine-check}
gradethis:: grade_code()

```
In the above graph, make it clearer that each line represents a different value for the drv variable by distinguishing each line with a different color as is done here:

```{r two-plus-nine2view}
ggplot(data = mpg) + geom_smooth(mapping = aes(x = displ, y = hwy,linetype = drv, color=drv))
```


```{r two-plus-nine2, exercise=TRUE}
ggplot(data = mpg) + 

```
```{r two-plus-nine2-hint-1}
"try geom _smooth() rather than geom_point()"
```

```{r two-plus-nine2-solution}
ggplot(data = mpg) + geom_smooth(mapping = aes(x = displ, y = hwy,linetype = drv, color=drv))
```

```{r two-plus-nine2-check}
gradethis:: grade_code()

```

### 3.6.1 Intermediate KA

<!-- ![picture](images/Rplot30.png) -->
<img src="images/Rplot30.png" alt="drawing" style="width:400px;"/>

create the above graph

```{r two-plus-ten, exercise=TRUE}

```

```{r two-plus-ten-hint-1}
"ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + "
```

```{r two-plus-ten-solution}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + 
  geom_point() + 
  geom_smooth()
```

```{r two-plus-ten-check}
gradethis::grade_code()

```

### 3.6.1 Intermediate YY

Change the variable that is assigned to linetype to "cyl" and predict what will be created before run the code. Below you can see the results of the current code.
```{r 336, exercise=TRUE, exercise.eval=TRUE}
ggplot(data = mpg) + 
  geom_smooth(mapping = aes(x = displ, y = hwy, linetype = drv))
```
```{r 336-hint-1}
"Be careful!! "
```
```{r 336-solution}
ggplot(data = mpg) + 
  geom_smooth(mapping = aes(x = displ, y = hwy, linetype = cyl))
```
```{r 336-check}
gradethis::grade_code()
```

Add the minimum code to allow this to run using displ for x and hwy for y

```{r 436, exercise=TRUE}
ggplot() + 
  geom_smooth()
```
```{r 436-hint-1}
"The function that causes the error was covered in section 3.5."
```
```{r 436-solution}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + 
  geom_smooth()
```
```{r 436-check}
gradethis::grade_code()
```




### Advanced 3.6.1 DN


Refer to the mpg dataset.  

Recall that 2-seater cars were somewhat of an anomoly with regard to the relationship between displacement and highway mileage. Write the R code to plot a geom_point graph together with geom_smooth graph, with color=class, excluding all 2seater cars  to generate this graph. Note you will need to use the filter() function.


```{r D24view}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + 
   geom_point(data = filter(mpg, class != "2seater"), mapping = aes(color = class)) + 
   geom_smooth(data = filter(mpg, class != "2seater"), se= FALSE)
```


```{r D24, exercise=TRUE}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + 
   geom_point(data =   ) + 
   geom_smooth(data = , se=FALSE)
```

```{r D24-hint}
"Actually you need to use the filter() function twice. Try just using it once then see if you can see where the second occurance should be. What happens when you only use it once?"
```

```{r D24-solution}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + 
   geom_point(data = filter(mpg, class != "2seater"), mapping = aes(color = class)) + 
   geom_smooth(data = filter(mpg, class != "2seater"), se= FALSE)
```

```{r D24-check}
gradethis::grade_this_code()
```

### 3.6.1 Challenging KA

<!-- <!-- ![picture](images/Rplot31.png) --> -->
<!-- <img src="images/Rplot31.png" alt="drawing" style="width:400px;"/> -->
```{r 361ChalKA}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + geom_point(mapping = aes(color = class)) + geom_smooth(data =mpg) 
```
finish the code to create the above graph

```{r two-plus-zero, exercise=TRUE}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + 
  geom_point(mapping = aes(color = class)) + 
  geom_smooth(
```



```{r two-plus-zero-solution}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + geom_point(mapping = aes(color = class)) + geom_smooth(data =mpg)  
```

```{r two-plus-zero-check}
gradethis:: grade_code()

```
```{r quiz4361KA}
question("Comparing the 2 above graphs, notice how the second one \"obscures\" the downward trend of displ and hwy. What class of cars is causing this to occur? ",
  answer("suv"),
  answer("minivan"),
  answer("2seater", correct = TRUE),
  answer("subcompact")
)

```
### 3.6.1 Challenging  YY

Create the graphs below.

<img src="https://d33wubrfki0l68.cloudfront.net/e4e1d3f9256c683b5219f9dd0eee81d0320e9664/cc55f/visualize_files/figure-html/unnamed-chunk-28-2.png" alt="drawing" style="width:400px;"/>


```{r 536, exercise=TRUE, exercise.eval=TRUE}

```
```{r 536-hint-1}
"use geom_point"
```
```{r 536-solution}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy)) +
  geom_smooth(mapping = aes(x = displ, y = hwy))
```
```{r 536-check}
gradethis::grade_code()
```
```{r quiz4361YY}
question("In the above graph the dots are probably too large. Where do we adjust the size of the points",
  answer("inside aes(); we place a third argument called size"),
  answer("inside geom_point(); we place a third argument called size", correct = TRUE),
  answer("inside ggplot(); we place a third argument called size"),
  answer("inside geom_smooth(); we place a third argument called size")
)

```

```{r quiz4361YY2}
question("Which is the best explanation for the previous answer?",
  answer("aes() is where we dictate what variables are mapped to which visual elements of the graph. The visual elements include the x-axes, the y-axis, the size of a point."),
  answer("aes() is where we dictate what variables are mapped to which visual elements of the graph. The visual elements include the x-axes, the y-axis, the size of a point. If we wanted the size of a point to be related to the value of a particular variable, then we would put size in aes()"),
  answer("aes() is where we dictate what variables are mapped to which visual elements of the graph. The visual elements include the x-axes, the y-axis, the size of a point. If we wanted the size of a point to be related to the value of a particular variable, then we would put size in aes(). If we want the size to be independent of any of our variables, then we would not put it in aes()."),
  answer("aes() is where we dictate what variables are mapped to which visual elements of the graph. The visual elements include the x-axes, the y-axis, the size of a point. If we wanted the size of a point to be related to the value of a particular variable, then we would put size in aes(). If we want the size to be independent of any of our variables, then we would not put it in aes(). We instead put it inside geom_point(). We can also control color for all points there, etc. ", correct = TRUE)
)

```
####################################################################


## Section 3.7.1

### 3.7.1 Easy YY


In this exercise we are going to use new dataset called diamonds.
Run the code to get to know diamonds. (note that View(diamonds) will show you the dataset but here we are thinking of another way to explore the dataset)
```{r 137, exercise=TRUE, exercise.eval=TRUE}

```
```{r 137-hint-1}
"Use ? "
```
```{r 137-solution}
?diamonds
```
```{r 137-check}
gradethis::grade_code()
```


```{r quiz37}
question("According to the text, which kind of graph computes a robust summary of the distribution and displays it as a specially formatted box?",
  answer("boxplots", correct = TRUE),
  answer("bar graph"),
  answer("scatter plot"),
  answer("spider chart")
)
```


####################################################################




### 3.7.1 Easy SD


```{r quiz2}
quiz(
  question(" ________ function is similar to the existing subset() function in base R but is quite a bit faster.",
    answer("rename"),
    answer("set"),
    answer("filter", correct = TRUE),
    answer("subset")
  )
)
```



```{r quiz3}
quiz(
  question(" Columns can be arranged in descending order  by using the special ____ operator.",
    answer("asc()"),
    answer("desc()", correct = TRUE),
    answer("descending()"),
    answer("subset")
  )
)
```


```{r quiz4}
quiz(
  question(" The ______ operator allows you to string operations in a left-to-right fashion.",
    answer("This %<%"),
    answer("This %%%"),
    answer("This %>%", correct = TRUE),
    answer("This >%>%")
  )
)
```


```{r quiz5}
quiz(
  question(" Which of the following returns a subset of the columns of a data frame?",
    answer("filter"),
    answer("get"),
    answer("select", correct = TRUE),
    answer("set"),
    answer("retrieve")
  )
)
```



### 3.7.1 Intermediate YY

There are two functions for creating barcharts that are quite similar: geom_bar() and geo_col(). Typically geo_bar() will work with just one column of data. For example, in the diamonds dataset, we could work with just the column called 'cut'. The levels of that column can be seen by running levels(diamonds$cut). This shows that there are 5 unique values: 
"Fair"      "Good"      "Very Good" "Premium"   "Ideal"
If you just want a bar chart that shows how many of each of these there are in the dataset you can use geom_bar.

```{r geo_bar1,  exercise=TRUE}
ggplot(data = diamonds,  ) +  
```
```{r geo_bar1-hint-1}
"Do not forget the aes() takes the variable of interest and we need geom_bar as well"
```
```{r geo_bar1-hint-2}
"ggplot(data = diamonds, aes(cut)) + geom_bar()"
```


```{r geo_bar1-solution}
ggplot(data = diamonds, aes(cut)) + geom_bar()
```



```{r geo_bar1-check}
gradethis::grade_code()
```

Notice we only used one column of data above (the cut column). If we want to use two columns of data; one for the x-axis and one for the heights of each of the bars above the x-axis, then we would use geom_col() instead.
Try running the following. You should check first to see that there are two columns of data being considered. In the code below, first we create a simplified dataset called df for this example. Second we create a graph. Note that df does have 2 columns and both will be used in the graph.

```{r geom_col1, exercise=TRUE, exercise.eval=TRUE}
df <- data.frame(trt = c("a", "b", "c"), outcome = c(2.3, 1.9,3.2))
ggplot(df, aes(trt, outcome)) +geom_col()

```



Run the code below. The code uses "stat_count". 
stat_count counts the number of cases at each x position. Change variable of x-axis to "color".
```{r 337, exercise=TRUE, exercise.eval=TRUE}
ggplot(data = diamonds) + 
  stat_count(mapping = aes(x = cut))
```
```{r 337-solution}
ggplot(data = diamonds) + 
  stat_count(mapping = aes(x = color))
```
```{r 337-check}
gradethis::grade_code()
```

<!-- #### Exercise 4 -->
<!-- ```{r quiz437} -->
<!-- question("what does prep in the code mean? ", -->
<!--   answer("variable prep"), -->
<!--   answer("count of each variable"), -->
<!--   answer("what percentage does each varlable occupy.", correct = TRUE), -->
<!--   answer("No answer here") -->
<!-- ) -->
<!-- ``` -->
### 3.7.1 Challenging  YY

Create the graphs below.(Top and bottom of each line and each point means max value, minimum value and median respectively.)

<img src="https://d33wubrfki0l68.cloudfront.net/4f7c8830283cb7009346ed85a8196b06959f5abd/78175/visualize_files/figure-html/unnamed-chunk-34-1.png" alt="drawing" style="width:400px;"/>

```{r 537, exercise=TRUE, exercise.eval=TRUE}

```
```{r 537-hint-1}
"the points on the lines are median. Thus, use fun = median"
```
```{r 537-solution}
ggplot(data = diamonds) + 
  stat_summary(
    mapping = aes(x = cut, y = depth),
    fun.min = min,
    fun.max = max,
    fun = median
  )
```
```{r 537-check}
gradethis::grade_code()
```

### 3.7.1 Challenging  SD

*Let's discover multiplot Can we draw it using R?*

For iris dataset, form a multiplot. Should we use facet_wrap? or facet_grid? Try to run the code and check to decide for yourself which is more informative.

```{r Multiplotgrid, exercise=TRUE}
#Facet_grid

iris
ggplot(data = iris,aes(x=Sepal.Length))+geom_histogram(binwidth = 0.1) + facet_grid(Species~.)
```

```{r multiplotwrap, exercise = TRUE}
library(tidyverse)
iris
ggplot(data = iris,aes(x=Sepal.Length))+geom_histogram(binwidth = 0.1) + facet_wrap(~Species)
```


## Section 5.2.4

### Easy 5.2.4 DN


```{r flights524dn}
head(flights)
 
colnames(flights)
```

Refer to flights dataset.
Find all flights that arrived on Christmas Eve (24th December)

```{r D5, exercise=TRUE}
filter(flights,     &     )
```

```{r D5-hint}
"Experiment with the filter() function"
```

```{r D5-solution}
filter(flights, month == 12 & day == 24)
```

```{r D5-check}
gradethis::grade_this_code()
```


### 5.2.4 Easy KA
```{r flights524KA}
head(flights)
 
colnames(flights)
```
Write the code to find the flights that had an arrival delay of 60 minutes or more.

```{r two-plus-8a, exercise=TRUE}

```
```{r two-plus-8a-hint-1}
"filter(flights, arr_delay >="
```

```{r two-plus-8a-solution}
filter(flights, arr_delay >= 60)
```
```{r two-plus-8a-check}
gradethis::grade_code()
```


### 5.2.4 Easy YY


In this exercise, we are going to use filter(). Run the code below (without the quotation marks) to explore the filter function.

```{r 152, exercise=TRUE}
"?filter()"
```


### 5.2.4 Easy SD

Create a new variable called delay which comes from the sum of arr_delay and dep_delay (you could use mutate for this but for this problem use flights$delay). In the 2nd line use summarise to find the mean delay (calling it mean_delay) and remove NAs from the delay using na.rm.


```{r q, exercise=TRUE}

```

```{r q-hint}
summarise(mean_delay = mean(delay,na.rm = TRUE))
```

```{r q-solution}
flights$delay <- flights$arr_delay + flights$dep_delay
summarise(flights, mean_delay = mean(delay, na.rm = TRUE))
```

```{r q-check}
gradethis::grade_this_code()
```




### Intermediate 5.2.4 DN  


Refer to flights dataset.
Find all flights with missing dep_delay, and arr_delay  

```{r D16, exercise=TRUE}
filter(flights,     &     )
```


```{r D16-hint}
"Experiment with is.na() and filter() function"
```

```{r D16-solution}
ggplot() +
  geom_point(data=mpg, mapping= aes(x=displ, y= hwy, color= drv)) +
  geom_line(data=mpg, mapping = aes(x=displ, y= hwy, color= drv))
```

```{r D16-check}
gradethis::grade_this_code()
```

### 5.2.4 Intermediate KA

Write the code to find flights that departed between January and August

```{r two-plus-threea, exercise=TRUE}

```
```{r two-plus-threea-hint-1}
filter(flights, month >= & month <= )

```
```{r two-plus-threea-solution}
filter(flights, month >= 7 & month <= 9)
```
```{r two-plus-threea-check}
gradethis::grade_code()
```

### 5.2.4 Intermediate YY


Find all flights that depart from LGA and EWR
```{r 352, exercise=TRUE, exercise.eval=TRUE}

```
```{r 352-hint-1}
"Use filter and | (or)"
```
```{r 523-solution}
filter(flights, dest == 'LGA' | dest == 'EWR')
```
```{r 352-check}
gradethis::grade_code()
```



### 5.2.4 Intermediate SD

*These questions will be in parts. Use 3 lines to write the answer*
Part1: Arrange flights in the order of year, month and day
Part2: Filter Delta flights departed in the month of November and December (carrier is DL)
How many rows are there now? (you can check flights -don't forget to remove the quotes)

```{r flights_e, exercise=TRUE, exercise.eval=TRUE}

```

```{r flights_e-hint}
# arrange(), inside the bracket there could be flights, year, month and day
# filter(flights, month %in% c(11, 12) & carrier == "DL")
nrow(filter(flights, month %in% c(11, 12) & carrier == "DL"))
```

```{r flights_e-solution}
arrange(flights,year, month, day)
filter(flights, month %in% c(11, 12) & carrier == "DL")
nrow(filter(flights, month %in% c(11, 12) & carrier == "DL"))
```

```{r flights_e-check}
gradethis::grade_this_code()
```


### Advanced 5.2.4 DN  


Refer to flights dataset.
Find all flights that belong to Mesa Airlines, Hawaiian Airlines, and Alaska Airlines

```{r D25-hint}
"check unique(flights$carrier)"
```

```{r D25, exercise=TRUE}
filter(flights, carrier ==    )
```

```{r D25-solution}
filter(flights, carrier == c("YV","HA","AS"))
```

```{r D25-check}
gradethis::grade_this_code(correct = 'Good work! You could use either filter(flights, carrier == c("YV","HA","AS"))   or filter(flights, carrier == "YV"| carrier == "HA"| carrier == "AS" ) ', incorrect = 'Not quite. {code_feedback()} {random_encouragement()}')
```  

### 5.2.4  KA

find the flights that Departed between midnight and 6am (inclusive). Note that midnight is coded as 2400 in dep_time.

```{r two-plus-fourb, exercise=TRUE}



```
```{r two-plus-fourb-solution}

filter(flights, dep_time <= 600 | dep_time == 2400)

```

```{r two-plus-fourb-check}
gradethis::grade_code()
```

### 5.2.4 Challenging  YY


Identify the same set of flights that code below selects but using R's not operator: !
```{r 552, exercise=TRUE, exercise.eval=TRUE}
filter(flights, arr_delay > 60, dep_delay > 60)
```
```{r 552-hint-1}
"!() means out of the specified range"
```
```{r 552-solution}
filter(flights, !(arr_delay <= 60 | dep_delay <= 60))
```
```{r 552-check}
gradethis::grade_code()
```
### 5.2.4 Challenging  SD

Missing data in the flights data set would suggest the flight being canceled so we would prefer to remove those. Remove all the rows if there is NA in arr_delay and dep_delay. Use the pipe operator: %>% 

```{r Remove_miss, exercise=TRUE, exercise.eval=TRUE}

```

```{r Remove_miss-solution}
flights %>% filter(!is.na(arr_delay) & !is.na(dep_delay))
```

```{r Remove_miss-check}
gradethis::grade_this_code()
```


####################################################################

#############################################
## Section 5.3.1 

### Easy 5.3.1 DN


```{r D6}
quiz(
  question("Which of the following are parameters in arrange() ?",
    answer("data"),
    answer(".data", correct= TRUE),
    answer("by_group"),
    answer(".by_group", correct= TRUE)
  )
)
```



```{r D7}
quiz(
  question("Using arrange() function, where are NA values located ?",
    answer("top of the data"),
    answer("randomized"),
    answer("end of the data", correct= TRUE)
  )
)
```


### 5.3.1 Easy KA

Sort flights to find the most delayed flights. 

```{r two-plus-1a, exercise=TRUE}

```
```{r two-plus-1a-hint-1}
"use desc( ))"
```
```{r two-plus-1a-solution}
arrange(flights, desc(dep_delay))
```
```{r two-plus-1a-check}
gradethis::grade_code()
```



### 5.3.1 Easy YY


Run the code below and figure out what is printed.
```{r 153, exercise=TRUE, exercise.eval=TRUE}
arrange(flights, dep_delay)
```


Use desc() to re-order by a column in descending order.
```{r 253, exercise=TRUE, exercise.eval=TRUE}
arrange(flights, dep_delay)
```
```{r 253-solution}
arrange(flights, desc(dep_delay))
```
```{r 253-check}
gradethis::grade_code()
```


### 5.3.1 Easy SD

```{r quiz_e}
quiz(
  question(" Columns can be arranged in descending order too by using the special ____ operator.",
    answer("asc()"),
    answer("desc()", correct = TRUE),
    answer("descending()"),
    answer("subset")
  )
)
```



### Intermediate 5.3.1 DN  


Refer to flights dataset.
Arrange flights based on descending dep_delay  

```{r D17, exercise=TRUE}
arrange(  ,   )
```

```{r D17-hint}
"Experiment with arrange() function"
```

```{r D17-solution}
arrange(flights, desc(dep_delay))
```

```{r D17-check}
gradethis::grade_this_code()
```


Refer to flights dataset.
Arrange flights based on ascending distance

```{r D18, exercise=TRUE}
arrange(    ,    )
```

```{r D18-hint}
"Ascending order is the default for arrange()"
```

```{r D18-solution}
arrange(flights, distance)
```

```{r D18-check}
gradethis::grade_this_code()
```


### 5.3.1 Intermediate KA

write the code to find the Flights that left the earliest.

```{r two-plus-2a, exercise=TRUE}

```

```{r two-plus-2a-solution}
arrange(flights, dep_delay)
```
```{r two-plus-2a-check}
gradethis::grade_code()
```



### 5.3.1 Intermediate YY


Find the fastest flight on March 5 2013.
```{r 353, exercise=TRUE, exercise.eval=TRUE}

```
```{r 353-hint-1}
"You can use asc()"
```


```{r quiz153}
question("Missing values always appear where in a sort? (you can use head() and tail() to check) ",
  answer("top"),
  answer("middle"),
  answer("anywhere you designate"),
  answer("end", correct = TRUE)
)
```


### 5.3.1 Intermediate SD

Although not particularly meaningful,make a new variable called delay which will be the total delay from departure to arrival.

```{r delay, exercise=TRUE, exercise.eval=FALSE}
flights$delay <-
```

```{r delay-hint}
flights$delay <- flights$arr_delay + flights$dep_delay
```

```{r delay-solution}
flights$delay <- flights$arr_delay + flights$dep_delay
```

```{r delay-check}
gradethis::grade_this_code(correct = "Good work! Shortly we will learn an alternative way to construct a new variable with the mutate() function", incorrect = "Not quite. {code_feedback()} {random_encouragement()}")
```




### Advanced 5.3.1 DN


### 5.3.1 Challenging KA

Use arrange() to show top 10 flights that traveled the farthest and top 10 that traveled the shortest? Use 2 lines of code.

```{r two-plus-3a, exercise=TRUE}
arrange()
arrange()
```


```{r two-plus-3a-solution}
arrange(flights, desc(distance))
arrange(flights, distance)
```
```{r two-plus-3a-check}
gradethis::grade_code()
```



### 5.3.1 Challenging  SD

Missing data in the flights dataset would suggest a flight being canceled so we may prefer to remove those. Remove all the rows if there is NA in arr_delay and dep_delay.

```{r Remove_missing, exercise=TRUE, exercise.eval=TRUE}

```

```{r Remove_missing-hint}
# is.na shows missing whereas !is.na shows not missing
```

```{r Remove_missing-solution}
flights %>% filter(!is.na(arr_delay) & !is.na(dep_delay))
```

```{r Remove_missing-check}
gradethis::grade_this_code()
```


#######################################


## Section 5.4.1

### Easy 5.4.1 DN  


Refer to flights dataset.
Rename the air_time column to airtime 

```{r D8, exercise=TRUE}
rename(flights,   =     )
```

```{r D8-hint}
"Experiment with the rename() function"
```


```{r D8-solution}
rename(flights, airtime = air_time)
```

```{r D8-check}
gradethis::grade_this_code()
```


### 5.4.1 Easy KA

Typically any_of(var) is used instead of all_of(var) when the set of variables (referred to as var here) may not be a subset of the variables in the tibble. However here we will just practice using the any_of() function to find dep_delay and dep_time vectors. Actually though it is not necessary and you could get the same result with select(flights, c())

```{r two-plus-1b, exercise=TRUE}

```

```{r two-plus-1b-hint-1}
"select(flights, any_of(c()))"
```

```{r two-plus-1b-solution}
select(flights, any_of(c("dep_time", "dep_delay")))
```

```{r two-plus-1b-check}
gradethis::grade_code()
```


### 5.4.1 Easy YY

Run ?select() to understand it.
```{r 154, exercise=TRUE, exercise.eval=TRUE}

```


### 5.4.1 Easy SD

```{r quiz541sd}
quiz(
  question(" The ______ operator allows you to string operations in a left-to-right fashion.",
    answer("This %<%"),
    answer("This %%%"),
    answer("This %>%", correct = TRUE),
    answer("This >%>%")
  )
)
```




### Intermediate 5.4.1 DN  


Refer to flights dataset.
Select all columns except year, month, day column  

```{r D19, exercise=TRUE}
select(flights,    )
```

```{r D19-hint}
"You can either list all other columns beside year, month, day but here we want to use `!` "
```

```{r D19-solution}
select(flights, !c(year, month, day))
```

```{r D19-check}
gradethis::grade_this_code()
```  

### 5.4.1 Intermediate KA

Get the tibble that contains dep_time, dep_delay, arr_time, arr_delay without using dep_time, dep_delay, arr_time, arr_delay in the function but rather use the starts_with() function.

```{r two-plus-2b, exercise=TRUE}

```
```{r two-plus-2b-solution}
select(flights, starts_with("dep_"), starts_with("arr_"))
```
```{r two-plus-2b-check}
gradethis::grade_code()
```



### 5.4.1 Intermediate SD

In a previous question, we created a new variable called delay using this code: flights\$delay <- flights\$arr_delay + flights\$dep_delay Now create this new variable again and then arrange it in descending order. Use two lines of code.
```{r arrange, exercise=TRUE, exercise.eval=TRUE}

```

```{r arrange-hint}
flights$delay <- flights$arr_delay + flights$dep_delay
arrange(flights, desc(delay))
```

```{r arrange-solution}
arrange(flights, desc(delay))
```

```{r arrange-check}
gradethis::grade_this_code()
```




### Advanced 5.4.1 DN  

Refer to flights dataset.
Select all columns not containing the word "time"

```{r D26, exercise=TRUE}
select(flights,    )
```

```{r D26-hint}
"Experiment with contains() function in dplyr package "
```

```{r D26-solution}
select(flights, !contains("time"))
```

```{r D26-check}
gradethis::grade_this_code()
```


### 5.4.1 Challenging KA

use the !! operator in a function to find the vectors of "dep_time", "dep_delay", "arr_time", "arr_delay".

```{r two-plus-3b, exercise=TRUE}

```
```{r two-plus-3b-hint-1}
 "Specify the names of the variables with a character vector and use the bang-bang operator (use c as a name for the variables) (!!)."
```
```{r two-plus-3b-solution}
variables <- c("dep_time", "dep_delay", "arr_time", "arr_delay")
select(flights, !!variables)
```
```{r two-plus-3b-check}
gradethis::grade_code()
```



### 5.4.1 Challenging  YY


Select the variables called "day" and "dep_delay" using starts_with(), ends_with(), contains().
The argument of each function must be one letter

```{r 554, exercise=TRUE, exercise.eval=TRUE}

```
```{r 5545-hint-1}
"use ? if you are not sure about these helper functions"
```

### 5.4.1 Challenging  SD

Find the 5 flights with the shortest time taken.

```{r question, exercise=TRUE, exercise.eval=TRUE}

```




########################################################


## Section 5.5.2

### Easy 5.5.2 DN



Experiment with these to understand the parameters .before and .after

```{r mutateLocation, exercise=TRUE}
df <- tibble(x = 1, y = 2)
df %>% mutate(z = x + y)
df %>% mutate(z = x + y, .before = 2)
df %>% mutate(z = x + y, .after = x)
```


```{r D9}
quiz(
  question("Which of the following arguments control where new columns should appear in mutate() ?",
    answer(".data"),
    answer(".keep"),
    answer(".before",  correct= TRUE),
    answer(".after", correct= TRUE)
  )
)
```


```{r D10}
quiz(
  question("Using mutate(), if a column is set to NULL, what will happen to it, for example with 'year' in ``flights %>% mutate(gain=arr_delay-dep_delay,year=NULL)`` ?",
    answer("removed", correct= TRUE),
    answer("retained"),
    answer("changed to NA"),
    answer("changed to logical(0)")
  )
)
```
```{r w1}
quiz(
  question("flights has 19 columns. If we run ``flights2 <- flights %>% mutate(gain=arr_delay-dep_delay,year=NULL)`` how many columns will flights2 have?",
    answer("18"),
    answer("20"),
    answer("19", correct= TRUE),
    answer("1")
  )
)

```

In preparation for the next question, obtain the column names for the flights dataset.
```{r w1_5, exercise=TRUE }
 
```

```{r w1_5-hint}
#use colnames
```
```{r w1_5-solution}
colnames(flights)
```
```{r w1_5-check}
gradethis::grade_this_code()
```


```{r w2}

quiz(
  question("To move the 'gain' column to be the first column, to what value could we set  the '.before' parameter in:  ``flights2 <- flights %>% mutate(gain=arr_delay-dep_delay, .before= )``",
    answer("0"),
    answer("1", correct= TRUE),
    answer("year", correct= TRUE),
    answer("any of these")
  )
)

```


Make a new column in the flights dataset called speed where speed = distance/air_time. Arrange flights based on descending speed
```{r D27, exercise=TRUE}
flights <- flights %>%
  mutate(speed =   /   )
  arrange(  ,  )
```

```{r D27-hint}
"Experiment with mutate() function in dplyr package "
```

```{r D27-solution}
flights <- flights %>%
  mutate(speed = distance/air_time)
  arrange(flights, desc(speed))
```

```{r D27-check}
gradethis::grade_this_code()
```


### 5.5.2 Easy KA

```{r quiz2b}
question("what would the following code ``1504 %/% 100`` give. Note to get help with these 'oddly named'  operators and functions you need to put backtics around their name like this ``?`%/%` `` ",
  answer("[1] 15", correct = TRUE),
  answer("[1] 4"),
  answer("150,400"),
  answer("15.04")
)
```



### 5.5.2  Easy YY


Run the code below and figure out what is mutate(). Don't forget to remove the quotation marks.
```{r 155, exercise=TRUE, exercise.eval=TRUE}
"?mutate"
```



```{r quiz155}
question("There are two functions that are useful for time based data: lag() and lead(). They are useful for comparing values behind or ahead of the current values, lag() pushing the data \"back\" and lead() pushing it \"forward\". If x <- 2:7, what is lag(x)?",
  answer("[1] NA  2  3  4  5  6", correct = TRUE),
  answer("[1] 2  3  4  5  6  NA"),
  answer("[1] 1  2  3  4  5  6"),
  answer("[1] 2  3  4  5  6  7")
)
```


### 5.5.2  Easy SD

Make a new variable called speed using the mutate function.
(Speed = Distance/ Time)
```{r speed,exercise=TRUE, exercise.eval=TRUE}

```


### Intermediate 5.5.2 DN



Refer to flights dataset.
Add a new column called air_time_hours by converting air_time from minutes to hours

```{r D20, exercise=TRUE}
mutate(    ,    )
```

```{r D20-hint}
"Experiment with mutate() function "
```

```{r D20-solution}
mutate(flights, air_time_hours = air_time/60)
```

```{r D20-check}
gradethis::grade_this_code()
```


Refer to flights dataset.
Add a new column called air_time_hours by converting air_time from minutes to hours but this time round *down* to the nearest hour.

```{r w3, exercise=TRUE}
mutate(    ,    )
```

```{r w3-hint}
"Experiment with mutate() function "
```

```{r w3-solution}
mutate(flights, air_time_hours = air_time%/%60)
```

```{r w3-check}
gradethis::grade_this_code()
```

Refer to flights dataset.
Add a new column called distance_km by converting distance from miles to km (1 mile = 1.61 km)  

```{r D21, exercise=TRUE}
mutate(    ,    )
```

```{r D21-hint}
"Experiment with mutate() function "
```

```{r D21-solution}
mutate(flights, distance_km = distance * 1.61)
```

```{r D21-check}
gradethis::grade_this_code()
```


### 5.5.2  Intermediate KA

```{r quiz1c}
question(" what does the following give us 1:3 + 1:10",
  answer("2  4  6  5  7  9  8 10 12 11", correct = TRUE),
  answer("2  4  6  5  7  9  9 10 12 11"),
  answer("2  3  4  5  6  7  8 9 10 11"),
  answer("gives us nothing")
)
```


### 5.5.2  Intermediate YY


Add new varible named "gain"
gain = dep_delay - arr_delay
```{r 355, exercise=TRUE, exercise.eval=TRUE}

```
```{r 355-hint-1}
"Use mutate()"
```
```{r 355-solution}
mutate(flights,
  gain = dep_delay - arr_delay,)
```
```{r 355-check}
gradethis::grade_code()
```


```{r 2quiz55}
question("Choose the best answer. transmute",
  answer("changes the content of the variable"),
  answer("create a new variable and deletes the old one", correct = TRUE),
  answer("changes the name of variable"),
  answer("does nothing")
)
```
### 5.5.2  Intermediate SD

This questions will be in parts. Use two lines of code to write your answer.
Part1: Make a new variable called delay which will be the total delay from departure to arrival.
Part2: Since we donot want some variables like arr_time, dep_time etc., let's drop it. (you can check flights -don't forget to remove the quotes)

```{r flights, exercise=TRUE, exercise.eval=TRUE}

```

```{r flights-hint}
# select() everything except arr_delay, dep_delay by putting a `` - `` before the vector
```
```{r flights-solution}
flights$delay <- flights$arr_delay + flights$dep_delay
select(flights, -c(arr_delay,dep_delay))
```

```{r flights-check}
gradethis::grade_this_code()
```


### Advanced 5.5.2 DN  


Refer to flights dataset.
Rank the 10 fastest flights and 10 slowest flights  

```{r D28-hint}
"Experiment with top_n() function"
```

```{r D28, exercise=TRUE}
flights <- mutate(flights, speed= distance/air_time)
arrange(  ,  ) %>% 
arrange(  ,  ) %>% 
```

```{r D28-solution}
flights <- mutate(flights, speed=distance/air_time)
arrange(flights, desc(speed)) %>% top_n(10)
arrange(flights, speed) %>% top_n(10)
```

```{r D28-check}
gradethis::grade_this_code()
```

### 5.5.2 Challenging KA

Compare dep_time, sched_dep_time, and dep_delay.To check this relationship, you will first need to convert the times to a form more amenable to arithmetic operations. for example, dep_time = (dep_time %/% 100 * 60 + dep_time %% 100) %% 1440,
```{r two-plus-1d2, exercise=TRUE}

```

```{r two-plus-1d2-hint-1}
flights_deptime <-
  mutate(flights,
    dep_time_min = (dep_time %/% 100 * 60 + dep_time %% 100) %% 1440,
    sched_dep_time_min = (sched_dep_time %/% 100 * 60 +
      sched_dep_time %% 100) %% 1440,
    dep_delay_diff = dep_delay - dep_time_min + sched_dep_time_min
  )
```

```{r two-plus-1d2-hint-2}
filter(flights_deptime, dep_delay_diff != 0)
```

```{r two-plus-1d2-solution}

filter(flights_deptime, dep_delay_diff != 0)
```

```{r two-plus-1d2-check}
gradethis::grade_code()
```


### 5.5.2  Challenging  YY


Order "flights" by year,month and day.
Then, create new variable that represents cumulative sum of  distance.

```{r 5, exercise=TRUE, exercise.eval=TRUE}

```
```{r 555-hint-1}
"Use arrange() to order the dataframe"
```

### 5.5.2  Challenging  SD

Now that we got the total delay, let's try to find if there is any relation between the distance travelled and delay. If there is any relation then analyze the case for each flight (carrier) using facet.

<img src="images/graph.png">


```{r arrange_delay, exercise=TRUE, exercise.eval=TRUE}
flights
```




###############################################
## Section 5.6.7

### Easy 5.6.7 DN  


Refer to flights dataset.
Find the mean air time for each month

```{r D11, exercise = TRUE}
flights %>% 
  group_by( ) %>%
  summarise(mean_air_time = mean( ))
```

```{r D11-hint}
"NAs should be removed in calculating the mean air time"
```

```{r D11-solution}
flights %>% 
  group_by(year, month) %>%
  summarise(mean_air_time = mean(air_time, na.rm = TRUE))
```

```{r D11-check}
gradethis::grade_this_code()
```


### 5.6.7 Easy KA

```{r quiz1e}
question("which delay may prove costlier to customer?",
  answer("arrival delay", correct = TRUE),
  answer("departure delay")
)
```


### 5.6.7  Easy YY

In this execise, we will use sumarize(). Run the code below to leanrn summrize() (remove the quotation marks)
```{r 1, exercise=TRUE, exercise.eval=TRUE}
"?summarise()"
```


### 5.6.7  Easy SD



In the flights dataset, correct the following code so as to find the flights departed in March and April (why is it wrong?): 

```{r Correct_the_code, exercise=TRUE}

filter(flights, month == (3 |4 ))
```

```{r Correct_the_code-hint}
# Use infix  (%in% or month == 3| month == 4)
 
```
```{r Correct_the_code-solution}
filter(flights, month == 3| month == 4)
```

```{r Correct_the_code-check}
gradethis::grade_this_code()
```




### Intermediate 5.6.7 DN  


Refer to flights dataset.
Find the median air time for each day of the year

```{r D22, exercise=TRUE}
flights %>% 
  group_by(   ) %>%
  summarise(median_air_time = median(   ))
```

```{r D22-hint}
"NAs should be removed in calculating the mean air time "
```

```{r D22-solution}
flights %>% 
  group_by(year, month, day) %>%
  summarise(median_air_time = median(air_time, na.rm = TRUE))
```

```{r D22-check}
gradethis::grade_this_code()
```

### 5.6.7 Intermediate KA


Give the code to show the mean delays (arr_delay) per carrier? Call it averageDelay.

```{r two-plus-foura, exercise=TRUE}

```
```{r two-plus-foura-hint-1}
"flights %>% group_by() %>% summarise()"

```
```{r two-plus-foura-solution}
flights %>%
  group_by(carrier) %>%
  summarise(averageDelay = mean(arr_delay, na.rm = TRUE))
```
```{r two-plus-foura-check}
gradethis::grade_code()
```

Give the code to show Which carrier has the worst delays (arr_delay)?

```{r two-plus-fourc, exercise=TRUE}

```
```{r two-plus-fourc-hint-1}
"sort the results"
```
```{r two-plus-fourc-solution}
flights %>%
  group_by(carrier) %>%
  summarise(averageDelay = mean(arr_delay, na.rm = TRUE)) %>%
  arrange(desc(averageDelay))

```
```{r two-plus-fourc-check}
gradethis::grade_code()
```




### 5.6.7  Intermediate YY


```{r quiz}
question("summarise() needs another function to be available. what is it?",
  answer("group_by()", correct = TRUE),
  answer("sort()"),
  answer("mutate()"),
  answer("arrange()")
)
```

### 5.6.7  Intermediate SD

Write your answer in two lines.
Line 1: Filter the Delta flights ("DL") for the month November and December and assign it a new variable named filtered_data .
Line 2: Show a tibble that has air_time, hour (the number of hours the flight took -rounded to the nearest hour), and minute (the total number of minutes the flight took  -the air_time converted to minutes)

```{r ranking, exercise=TRUE, exercise.eval=TRUE}
 #filtered_data <-
```

```{r ranking-hint1}
# First filter DL flights in only the last month using filter(). Don't forget to use infix %in%
# Save it to filtered_data
# Then use transmute() where hour = air_time %/% 60 and minute is air_time - hour*60
```

```{r ranking-hint2}
transmute(flights,
  air_time,
  hour = air_time %/% 60,
  minute = air_time - hour*60
)
```

```{r ranking-solution}
filtered_data <- filter(flights, month %in% c(11, 12) & carrier == "DL")
transmute(flights,
  air_time,
  hour = air_time %/% 60,
  minute = air_time - hour*60
)
```

```{r ranking-check}
gradethis::grade_this_code()
```


### Advanced 5.6.7 DN  


Refer to flights dataset.
Find the first departure and last arrival for each day of the year, and plot the following graph  

```{r D29-graph}
flights %>%
  filter(!is.na(dep_time)&!is.na(arr_time)) %>%
  group_by(year, month, day) %>%
  summarise(first_dep = first(dep_time),
            last_arr = last(arr_time)) %>%
  ggplot(aes(first_dep,last_arr)) +
  geom_smooth()
```

```{r D29, exercise=TRUE}
flights %>%
  filter(    &    ) %>%
  group_by(    ) %>%
  summarise(first_dep = first(   ),
            last_arr  =  last(   )) %>%
  ggplot(aes(first_dep,last_arr)) +
  geom_smooth()
```

```{r D29-hint}
"Refer back to prevous exercises which use is.na(), group_by(), and summarise()"
```

```{r D29-solution}
flights %>%
  filter(!is.na(dep_time)&!is.na(arr_time)) %>%
  group_by(year, month, day) %>%
  summarise(first_dep = first(dep_time),
            last_arr = last(arr_time)) %>%
  ggplot(aes(first_dep,last_arr)) +
  geom_smooth()
```

```{r D29-check}
gradethis::grade_this_code()
```

### 5.6.7 Challenging KA

What time of day (what hour) should you fly if you want to avoid delays as much as possible?
Give the code to generate a tibble whose first column is the hour, and whose second column shows the mean of the arrival delays (call it meanDelay) and display it sorted from least to greatest. 
 
 ```{r two-plus-fivec, exercise=TRUE}

```
```{r two-plus-fivec-solution}
flights %>%
  group_by(hour) %>%
  summarise(meanDelay = mean(arr_delay, na.rm = TRUE)) %>%
  arrange(meanDelay)
```
```{r two-plus-fivec-check}
gradethis::grade_code()
```

```{r quizKABestTime}
question("Which hour of the day has the least average arrival delays?",
  answer("10"),
  answer("12"),
  answer("16"),
  answer("7", correct = TRUE),
  answer("5")
)
```

### 5.6.7  Challenging  YY

The code below attempts to obtain the mean for arr_delay per year, month and day.
However it is not working. 
Modify the code to show the mean of arr_delay per year, month and day.
```{r 556, exercise=TRUE, exercise.eval=TRUE}
not_cancelled <- flights %>% 
  filter(!is.na(dep_delay), !is.na(arr_delay))
not_cancelled %>% 
  group_by(year, month, day) %>% 
  summarise(
    meanDelayPerDay = sum(arr_delay)
  )
```
```{r 556-solution}
not_cancelled <- flights %>% 
  filter(!is.na(dep_delay), !is.na(arr_delay))
not_cancelled %>% 
  group_by(year, month, day) %>% 
  summarise(
    meanDelayPerDay = mean(arr_delay)
  )
```
```{r 556-check}
gradethis::grade_code()
```


### 5.6.7  Challenging  SD

Let's remove all the variables whose name contains the string 'time' in them with the contains() function. 

```{r Remove, exercise=TRUE, exercise.eval=TRUE}

```

```{r Remove-hint}
select(flights, -contains("time"))
```

```{r Remove-solution}

select(flights, -contains("time"))
```

```{r Remove-check}
gradethis::grade_this_code(correct = "You nailed it! You may want to think about whether we can use any_of() to do the same.", incorrect = "Not quite. {code_feedback()} {random_encouragement()}")
```
## Section 5.7.1   

### 5.7.1   Intermediate SD

Find the mean delay of flights given that arr_delay > 30 minutes and among those flights find the carrier with the highest delay.


```{r print-limit_s, exercise=TRUE, exercise.eval=TRUE}

```

```{r print-limit_s-hint}
flights %>%
	filter(arr_delay > ) %>%
	group_by() %>%
	summarise(mean_delay = mean()) %>%
  arrange(desc(mean_delay))
```

```{r print-limit_s-solution}
flights %>%
	filter(arr_delay > 30) %>%
	group_by(carrier) %>%
	summarise(mean_delay = mean(arr_delay)) %>%
  arrange(desc(mean_delay))
```

```{r print-limit_s-check}
gradethis::grade_this_code()
```

### 5.7.1   Challenging  SD

Using flights data, group data by carrier and summarize the data by arr_delay. Finally, plot the summary statistics like this. Note that we chose not to display the legend and the theme being used is theme_classic()

```{r summaryofdestinations}
flights %>% group_by(carrier) %>% summarise(mean_arr_delay = mean(arr_delay,na.rm = TRUE)) %>% ggplot(aes(x=carrier, y = mean_arr_delay, fill = carrier)) + geom_bar(stat = "identity",show.legend = FALSE) + theme_classic()
```


```{r summary_stats, exercise=TRUE, exercise.eval=TRUE}

#flights %>% group_by() %>% summarise(mean_arr_delay = mean(arr_delay,na.rm = TRUE)) %>%

```

```{r summary_stats-hint}
#flights %>% group_by(carrier) %>% summarise(mean_arr_delay = mean( ,na.rm = TRUE)) %>% ggplot(aes(x= , y = mean_arr_delay, fill = carrier)) + geom_bar(stat = "identity",show.legend = ) + theme_classic()
```

```{r summary_stats-solution}
flights %>% group_by(carrier) %>% summarise(mean_arr_delay = mean(arr_delay,na.rm = TRUE)) %>% ggplot(aes(x=carrier, y = mean_ardelay, fill = carrier)) + geom_bar(stat = "identity",show.legend = FALSE) + theme_classic()
```

```{r summary_stats-check}
gradethis::grade_this_code()
```


####################################







